---
title       : Insert the chapter title here
description : Insert the chapter description here
attachments :
  slides_link : https://s3.amazonaws.com/assets.datacamp.com/course/teach/slides_example.pdf


--- type:NormalExercise lang:r xp:100 skills:1 key:5347fc2fc3
## Go solve this wallnut!

Word games are easy...unless they are in Spanish!

*** =instructions
-  Entero: Le haces esto a un cadáver 
-  Decapitado: Es lo que haces en un bote

*** =hint
- Hints are for the weak

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}
#write answer to Entero in the form Entero="answer here, no caps" 
Entero=""
#write answer to Decapitado in the form Decapitado="answer here, no caps"
Decapitado=""
```

*** =solution
```{r}
#write answer to Entero in the form Entero="answer here, no caps" 
Entero="cremar"
#write answer to Decapitado in the form Decapitado="answer here, no caps"
Decapitado="remar"
```

*** =sct
```{r}
# SCT written with testwhat: https://github.com/datacamp/testwhat/wiki



test_object("Entero")
test_object("Decapitado")

test_error()

success_msg("Happy birthday little garbanzo!!")
```

--- type:NormalExercise lang:r xp:100 skills:1 key:5c168e917f
## But wait there's more!

Word games are easy...unless they are in Spanish!

*** =instructions
-  Take the ‘Entero’ word from the previous clue. Create a new word by adding one letter that, in one liddo dialect, is replaced by a ‘t’. Remove the last letter of that word and add a specific constant. This word is a place that was important to your time in San Francisco. 
-  Give your answer as one word, no caps, in the form Answer=""



*** =hint
- Peanuts need no hints

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}
#write your answer here
Answer=""
```

*** =solution
```{r}
#write your answer here
Answer="market"
```

*** =sct
```{r}
# SCT written with testwhat: https://github.com/datacamp/testwhat/wiki



test_object("Answer")


test_error()

success_msg("We couldn’t imagine a better person to have as a PSC member!")
```
