---
title       : Insert the chapter title here
description : Insert the chapter description here
attachments :
  slides_link : https://s3.amazonaws.com/assets.datacamp.com/course/teach/slides_example.pdf


--- type:NormalExercise lang:r xp:100 skills:1 key:5347fc2fc3
## Go solve this walnut!

Word games are easy...unless they are in Spanish!

*** =instructions
-  Entero: Le haces esto a un cadáver 
-  Decapitado: Es lo que haces en un bote

*** =hint
- Hints are for the weak

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}
#write answer to Entero in the form Entero="answer here, no caps" 
Entero=""
#write answer to Decapitado in the form Decapitado="answer here, no caps"
Decapitado=""
```

*** =solution
```{r}
#write answer to Entero in the form Entero="answer here, no caps" 
Entero="cremar"
#write answer to Decapitado in the form Decapitado="answer here, no caps"
Decapitado="remar"
```

*** =sct
```{r}
# SCT written with testwhat: https://github.com/datacamp/testwhat/wiki



test_object("Entero")
test_object("Decapitado")

test_error()

success_msg("Happy birthday little garbanzo!!")
```

--- type:NormalExercise lang:r xp:100 skills:1 key:5c168e917f
## But wait there's more!

Word games are easy...unless they are in Spanish!

*** =instructions
-  Take the ‘Entero’ word from the previous clue. Create a new word by adding one letter that, in one liddo dialect, is replaced by a ‘t’. Remove the last letter of that word and add a specific constant. This word is a place that was important to your time in San Francisco. 
-  Give your answer as one word, no caps, in the form Answer=""



*** =hint
- Peanuts need no hints

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}
#write your answer here
Answer=""
```

*** =solution
```{r}
#write your answer here
Answer="market"
```

*** =sct
```{r}
# SCT written with testwhat: https://github.com/datacamp/testwhat/wiki



test_object("Answer")


test_error()

success_msg("We couldn’t imagine a better person to have as a PSC member!")
```
--- type:NormalExercise lang:r xp:100 skills:1 key:5c168e917f
## Almost there!!

Or maybe not...you never know!

*** =instructions
-  Give numerical answers to the following




*** =hint
- Peanut, if you keep asking for hints, we'll turn you into peanut butter

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}
#the number of spellers in the 2016 bee who had relatives who had participated in a National Spelling Bees
B1=
#the spelling bee number of the only non-Klionsky who is a native one liddo dialect speaker
B2=
#the spelling bee number of the only 2010 speller coming from Puerto Rico
B3=
#the spelling bee number of the real BEE in the 2010 bee and the Akiba baby in the 2016 bee
B4=
```

*** =solution
```{r}
#the number of spellers in the 2016 bee who had relatives who had participated in a National Spelling Bees
B1=29
#the spelling bee number of the only non-Klionsky who is a native one liddo dialect speaker
B2=89
#the spelling bee number of the only 2010 speller coming from Puerto Rico
B3=214
#the spelling bee number of the real BEE in the 2010 bee and the Akiba baby in the 2016 bee
B4=59
```

*** =sct
```{r}
# SCT written with testwhat: https://github.com/datacamp/testwhat/wiki



test_object("B1")
test_object("B2")
test_object("B3")
test_object("B4")


test_error()

success_msg("Conveniently, all of these numbers multiplied together (32,587,706) probably approaches the number of minutes we have all spent talking, walking, laughing, etc. (fermi estimation, ok!) . Here’s to hundreds of thousands more.")
```

